---
title: Improved Quadratic Equation Solver
skip-execution: true
---

```python
import math
import numpy as np
from ipywidgets import interact
```

```python
if not input('Load JupyterAI? [Y/n]').lower()=='n':
    %reload_ext jupyter_ai
```

In this notebook, we will improve the quadratic equation solver in the previous lab using conditional executions.

## Discriminant

Recall that a quadratic equation is

$$
ax^2+bx+c=0
$$ (eq:imp:quadratic)

where $a$, $b$, and $c$ are real-valued coefficients, and $x$ is the unknown variable.

::::{prf:definition} Discriminant


The disciminant of a quadratic equation is defined as

$$
\Delta := b^2 - 4ac.
$$ (eq:discriminant)

::::

The discriminant $\Delta$ discriminates the roots 

$$\frac{-b\pm \sqrt{\Delta}}{2a}$$ (roots)

of a quadratic equation:

::::{prf:proposition}

The roots of a quadratic equation [](#eq:imp:quadratic) are both equal to 

$$-\frac{b}{2a}$$ (eq:repeated_root)

if and only if the discriminant is zero, i.e., $\Delta=0$.

::::

For example, if $(a,b,c)=(1,-2,1)$, then 

$$\Delta = (-2)^2 - 4(1\cdot 1) = 0$$

and so the repeated root is

$$
- \frac{b}{2a} = \frac{2}2 = 1.
$$

$(x-1)^2 = x^2 - 2x + 1$ as expected.

We will improve the quadratic solver to return only one root when the discriminant is close enough to $0$.

::::{exercise}
:label: ex:rel_tol

To determine whether a floating point number $x$ is close to $0$, the implementation `math.isclose` is insufficient:

- The test `math.isclose(x, 0, rel_tol=rel_tol)` is ineffective as explained in the lecture as it boils down to `x==0` or `rel_tol>=1`.
- The test `math.isclose(x, 0, abs_tol=1e-9)` may not work as expected since `x` (and therefore its precision error) floats at different scale according to its exponent.

Complete the following function to implement an alternative test of how close $x$ is to `0` relative to $y$ and $z$ with a relative tolerance of $\delta_{\text{rel}}$:

$$
\lvert x\rvert\leq  \max\Set{\lvert y\rvert, \lvert z\rvert} \delta_{\text{rel}}
$$ (eq:quad:rel_tol)

This test is similar to `math.isclose`, where $\delta_{\text{rel}}\geq 0$ is the relative tolerance specified by the [keyword-only argument](https://peps.python.org/pep-3102/) `rel_tol`, which has a default value of `1e-9`.

::::

<!-- grade_id: iszero, points: 0, grade: false, solution: true -->
```python
def iszero(x, y, z, *, rel_tol=1e-9):
    # YOUR CODE HERE
    raise NotImplementedError
```

<!-- grade_id: test_iszero, points: 1, grade: true, solution: false -->
```python
# tests
assert not iszero(1e-8, 0, 1)
assert iszero(1e-8, 1, 0, rel_tol=1e-7)
assert iszero(1e-8, 1, 10)
```

<!-- grade_id: htest_iszero, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

::::{exercise}  
:label: zero-determinant

Complete the following function by assigning to the variable`roots` a single root when the discriminant is close to zero relative to `b**2` and `4ac`, using the relative tolerence `rel_tol` that defaults to `1e-9`. For example, if `(a, b, c) == (1, -2, 1)`, then `roots` should be assigned the value `1.0` instead of `1.0, 1.0`.

::::

<!-- grade_id: zero_determininant, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c, *, rel_tol=1e-9):
    if iszero(d:=((y:=b**2)-(z:=4*a*c)), y, z, rel_tol=rel_tol):
    # YOUR CODE HERE
    raise NotImplementedError
    return roots
```

<!-- grade_id: test-zero_determinant, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
assert np.isclose(get_roots(1, 0, 0), 0.0).all()
assert np.isclose(get_roots(1, 1e-9, 0), 0.0).all()
assert np.isclose(get_roots(1, 1e-6, 0), (-1e-06, 0.0)).all()
```

<!-- grade_id: hidden_test-zero_determinant, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

::::{exercise}
:label: ex:isclose

To check whether the root is repeated, why is it better to use `iszero(d:=((y:=b**2)-(z:=4*a*c)), y, z, rel_tol=rel_tol)` instead of `d:=((y:=b**2)-(z:=4*a*c) == 0`?

::::

<!-- grade_id: isclose, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Linear equation

::::{exercise}
:label: ex:zero-a

Give the name of the error that the formula {eq}`roots` implemented in Python will raise when $a=0$. Explain whether it is a runtime error or not.

::::

<!-- grade_id: zero-a, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Nevertheless, the quadratic equation remains valid:

$$
bx + c=0,
$$

the root of which is $x = -\frac{c}b$.

::::{exercise}
:label: linear

Improve the function `get_roots` to return the root $-\frac{c}{b}$ when $a$ is close to zero relative to the other coefficients using the relative tolerance `rel_tol`.

::::

<!-- grade_id: linear, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c, *, rel_tol=1e-9):
    # YOUR CODE HERE
    raise NotImplementedError
    return roots
```

<!-- grade_id: test-linear, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
assert np.isclose(get_roots(1, 0, 0), 0.0).all()
assert np.isclose(get_roots(1, -2, 1), 1.0).all()
assert np.isclose(get_roots(0, -2, 1), 0.5).all()
```

<!-- grade_id: hidden_test-linear, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

## Degenerate cases

What if $a=b=0$? In this case, the equation becomes

$$
c = 0
$$
which is always satisfied if $c=0$ but never satisfied if $c\neq 0$.

::::{exercise}
:label: degenerate

Improve the function `get_roots` to return root(s) under all the following cases:
- If $a=0$ and $b\neq 0$, assign `roots` to the single root $-\frac{c}{b}$. 
- If $a=b=0$ and $c\neq 0$, assign `roots` to `None`.
- If $a=b=c=0$, there are infinitely many roots. Assign to `roots` the tuple `-float('inf'), float('inf')`.  
    Note that `float('inf')` converts the string `'inf'` to a floating point value that represents $\infty$.

The above checks of whether a coefficient is `0` should be relative to other coefficients with a relative tolerance of `rel_tol`.

:::{caution}
- `None` is a *Python keyword* that refers to a special object. You should *NOT* regard it as a string or quote it like `"None"`.
- `float('inf')` converts the string `'inf'` to a floating point number no smaller than any other floating point numbers. You should *NOT* write `inf`, which needs not refer to `float('inf')`.
:::

::::

<!-- grade_id: degenerate, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c, *, rel_tol=1e-9):
    # YOUR CODE HERE
    raise NotImplementedError
    return roots
```

<!-- grade_id: test-degenerate, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
assert np.isclose(get_roots(1, 0, 0), 0.0).all()
assert np.isclose(get_roots(1, -2, 1), 1.0).all()
assert np.isclose(get_roots(0, -2, 1), 0.5).all()
assert np.isclose(get_roots(0, 0, 0), (-float('inf'), float('inf'))).all()
assert get_roots(0, 0, 1) is None
```

<!-- grade_id: test_test-degenerate, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

After you have complete the exercises, you can run your robust solver below:

```python
# quadratic equations solver
@interact(a=(-10, 10, 1), b=(-10, 10, 1), c=(-10, 10, 1))
def quadratic_equation_solver(a=1, b=2, c=1):
    print("Root(s):", get_roots(a, b, c))
```